# Домашнее задание

В данном задании предлагается ознакомиться с форматом сдачи домашних заданий и решить одну задачу по теме if-else.

**1.** Напишите функцию, которая возвращает максимум из 3-х чисел. На вход функции подаются 3 числа, ожидаемые выход - одно максимальное число. **Использование max() запрещено.**

<!-- grade_id: cell-709514dd21c41c0f, points: 0, grade: false, solution: true -->
```python
def max_three(a, b, c):
    if a > b:
        if a > c:
            return a
        else:
            return c
    else:
        if b > c:
            return b
        else:
            return c
```

<!-- grade_id: cell-1fdb6a3ab3d61441, points: 2, grade: true, solution: false -->
```python
assert (max_three(1,2,3) == 3)
assert (max_three(100,100,150) == 150)
```

```python

```