# Домашнее задание

В данном задании предлагается решить несколько задач c помощью стандартных функций языка Python. Импортировать какие-либо библиотеки для выполнения заданий не требуется.
Формат входных и выходных данных будет описан в условиях каждой задачи.

**1.** Напишите функцию, которая меняет порядок слов в предложении на обратный. Вход строка из слов, выход: строка со словами в обратном порядке. Для упрощения на вход подается строка, которая не содержит знаков препинания.

<!-- grade_id: cell-4fd533c24fa05598, points: 0, grade: false, solution: true -->
```python
def reverse_word(s):
    words = s.split()
    words.reverse()
    return " ".join(words)
```

<!-- grade_id: cell-e459807e422a508c, points: 2, grade: true, solution: false -->
```python
assert (reverse_word('Мама мыла раму') == 'раму мыла Мама')
```

**2.** Напишите функцию, которая генерирует словарь в котором ключом является натуральное число $х$, а значением квадрат этого числа $x^2$. На вход функции подается число $n$ - количество пар в словаре, ключи словаря - последовательные натуральные число от 1 до $n$. Выход функции - словарь.

<!-- grade_id: cell-ead263b6212e872a, points: 0, grade: false, solution: true -->
```python
def dict_squares(n):
    result = {}
    for i in range(1, n + 1):
        result[i] = i * i
    return result
```

<!-- grade_id: cell-e7feadcff5905449, points: 2, grade: true, solution: false -->
```python
assert (dict_squares(3) == dict({1:1, 2:4, 3:9}))
```

**3.** Напишите функцию, которая принимает на вход натуральное число $n$. Функция возвращает список из натуральных делителей этого числа. Список отсортирован по возрастанию.

<!-- grade_id: cell-856052d6f10345b4, points: 0, grade: false, solution: true -->
```python
def divisors(n):
    result = []
    for i in range(1, n + 1):
        if(n % i == 0):
            result.append(i)
    return result
```

<!-- grade_id: cell-4fa27dc912c48192, points: 2, grade: true, solution: false -->
```python
assert (divisors(44) == [1,2,4,11,22,44])
assert (divisors(1) == [1])
```

**4.** (\*) Напишите функцию, которая считает количество вхождений подстроки в строку. На вход подается два аргумента: строка и подстрока, вхождение которой необходимо определить. На выход возвращается 1 число - количество вхождений. **Замечание**: подстроки могут накладываться друг на друга, к примеру для входа: строка - "ооо", подстрока - "оо" верный ответ 2.

<!-- grade_id: cell-1a54a1eeb74231ce, points: 0, grade: false, solution: true -->
```python
def pattern_occurrence(string, pattern):
    start = 0
    count = 0
    while(start < len(string)):
        pos = string.find(pattern, start)
        if(pos != -1):
            count = count + 1
            start = pos + 1
        else:
            break
        
    return count
```

<!-- grade_id: cell-1f5b848784dd3e1f, points: 2, grade: true, solution: false -->
```python
assert (pattern_occurrence('колокол', 'кол') == 2)
assert (pattern_occurrence('ooo', 'oo') == 2)
```

**5.** (\*) Реализуйте функцию, которая принимает один аргумент число $n: n > 0$ и возвращает $n$-ое число Фибоначчи. Последовательность чисел Фибоначчи определяется формулой $F_n = F_{n-1} + F_{n-2}$, $F_1 = 1, F_2 = 1$. Каждое следующее число является суммой двух предыдущих. Реализация может быть рекурсивной, но это не обязательное условие.

<!-- grade_id: cell-f0b8274144f1d5ad, points: 0, grade: false, solution: true -->
```python
def fib(n):
    if(n < 3):
        return 1
    else:
        return fib(n-1) + fib(n-2)
```

<!-- grade_id: cell-04838228a4238ad8, points: 2, grade: true, solution: false -->
```python
assert (fib(3) == 2)
assert (fib(10) == 55)
```

```python

```