# Домашнее задание


**1.** В папке с домашним заданием находится модуль mymodule, импортируйте его и вызовите функцию hello.

<!-- grade_id: cell-dc683677294cc842, points: 0, grade: false, solution: true -->
```python
import mymodule
mymodule.hello()
```

Output:
```
hello, world
```

**2.** Посчитайте синус угла 0.34 (в радианах), для этого можно воспользоваться библиотекой math. Импортируйте её и вызовите функцию для вычисления синуса. Присвойте значение синуса 0.34 переменной х.

```python
x = 0
```

<!-- grade_id: cell-24f5311033684282, points: 0, grade: false, solution: true -->
```python
import math
x = math.sin(0.34)
```

<!-- grade_id: cell-512ff4c3f4ffca65, points: 1, grade: true, solution: false -->
```python
assert (x > 0)
```

**3.** Создайте функцию is_prime(x), которая проверяет является ли натуральное число x простым. Возвращает True - если x простое, False - иначе.

Будьте внимательны! Подумайте о тех случаях, когда может возникнуть бесконечный цикл.

<!-- grade_id: cell-f4e253f23592c55f, points: 0, grade: false, solution: true -->
```python
import math
def is_prime(x):
    if(x <= 1):
        return False
    if(x <= 3):
        return True
    if(x % 2 == 0 or x % 3 == 0):
        return False
    for i in range(5, math.ceil(math.sqrt(x))):
        if(x % i == 0):
            return False
    return True
```

<!-- grade_id: cell-be6d8735cdb44437, points: 2, grade: true, solution: false -->
```python
# assert (is_prime(1) == False) # Перед тем, как раскомментировать и выполнить этот тест, убедитесь, что не возникнет бесконечного цикла!
assert (is_prime(2) == True)
assert (is_prime(4) == False)
assert (is_prime(9) == False)
assert (is_prime(11) == True)
```

**4.** Пользователь делает вклад в размере $x$ рублей сроком на $n$ лет под 8% годовых (каждый год размер его вклада увеличивается на 8%. Эти деньги прибавляются к сумме вклада, и на них в следующем году тоже будут проценты).

Написать функцию bank, принимающая аргументы $x$ и $n$, и возвращающую сумму, которая будет на счету пользователя.

<!-- grade_id: cell-9ee1275ffa6208c6, points: 0, grade: false, solution: true -->
```python
def bank(x, n):
    sum = x
    for i in range(n):
        sum = sum * 1.08
    return sum
```

<!-- grade_id: cell-d884635dcb461f8f, points: 2, grade: true, solution: false -->
```python
assert (bank(100, 1) == 108)
assert (bank(100, 10) - 215.8924997 < 0.0001)
```

**5.** Напишите функцию fio, которая преобразует строку, содержащую фамилию, имя и отчество человека (каждое слово с большой буквы), к форме <фамилия> <инициалы>.
Пример: 'Иванов Иван Иванович' -> 'Иванов И.И.'

<!-- grade_id: cell-dde3dbcf144ed6ee, points: 0, grade: false, solution: true -->
```python
def fio(s):
    terms = s.split()
    return terms[0] + " " + terms[1][0] + "." + terms[2][0] + "."
```

<!-- grade_id: cell-48b920c68cd2b765, points: 2, grade: true, solution: false -->
```python
assert(fio("Иванов Иван Иванович") == 'Иванов И.И.')
```

**6.** Напишите функцию mult. На вход подается строка, имеющая вид "$p_1$ * $p_2$ * ... * $p_n$", где $p_i$ – цифра. Вычислите значение выражения.

Пример: '2*3' -> 6

<!-- grade_id: cell-6e52737964f74e8a, points: 0, grade: false, solution: true -->
```python
def mult(s):
    terms = s.split("*")
    res = 1
    for n in terms:
        res = res * int(n)
    return res
```

<!-- grade_id: cell-703b95080ab5fb6d, points: 2, grade: true, solution: false -->
```python
assert(mult("5*2*3") == 30)
assert(mult("1*2*3") == 6)
```

**7.** Проверьте своё понимание областей видимости попробовав угадать, что выведет следующий код.

```python
a = 10
b = 20
c = 30
d = [40]
e = 50

def func(b=0):
        print(f'a = {a}')
        print(f'b = {b}')
        c = 7
        print(f'c = {c}')
        d.append(31)
        print(f'd = {d}')
        try:
            print(f'e = {e}')
            e += 8
        except Exception as ex:
            print(ex)

func()
```

Output:
```
a = 10
b = 0
c = 7
d = [40, 31]
local variable 'e' referenced before assignment
```

Ответьте на вопросы:

1. Что будет выведено на экран и почему, если после вызова функции `func()` выполнить `print(c)`?
2. Почему появилась ошибка? Исчезнет ли она, если `e += 8` заменить на `e = 8`? Исчезнет ли она, если `e += 8` выполнить до `print(f'e = {e}')`?  

<!-- grade_id: cell-9e3de693265b3248, points: 2, grade: true, solution: true -->
1.30
2. Попытка обращения к локальной переменной до ее объявления. Ошибка не исчезнет в обоих случаях, так как перед использованием локальную переменную нужно объявить, а e += 8 cразу же пытается ее читать.

```python
# Напоследок пример того как делать не надо.
# Не переопределяйте встроенные имена!
print = 5
print(print)
# После выполнения этой ячейки вы не сможете впоспользоваться функцией print.
# Чтобы перезапустить окружение нажмите в верхней панели Kernel->Restart.
# Весь код и выводы ячеек сохранятся.
```

Output:
```
TypeError                                 Traceback (most recent call last)
<ipython-input-14-d2f87c4faea4> in <module>
      2 # Не переопределяйте встроенные имена!
      3 print = 5
----> 4 print(print)
      5 # После выполнения этой ячейки вы не сможете впоспользоваться функцией print.
      6 # Чтобы перезапустить окружение нажмите в верхней панели Kernel->Restart.

TypeError: 'int' object is not callable
```